## HW01: DATA COLLLECTING

This is Assignment 01 for the course "Introduction to Data Science" at the Faculty of Information Technology, University of Science, Vietnam National University, Ho Chi Minh City.

(Latest update: 09/08/2024)

Student Name:Tong Trong Tam 

Student ID: 22120322

## **Assignment Objectives**

By completing this assignment, students will achieve the following objectives:
- Understand HTML Parsing
- Data Crawling via APIs
- Construct Dataframes
- Error Handling and Robustness

## **How to Complete and Submit the Assignment**

&#9889; **Note**: You should follow the instructions below. If anything is unclear, you need to contact the teaching assistant or instructor immediately for timely support.

**How to Do the Assignment**

You will work directly on this notebook file. First, fill in your full name and student ID (MSSV) in the header section of the file above. In the file, complete the tasks in sections marked:
```python
# YOUR CODE HERE
raise NotImplementedError()
```
Or for optional code sections:
```python
# YOUR CODE HERE (OPTION)
```
For markdown cells, complete the answer in the section marked:
```markdown
YOUR ANSWER HERE
```

**How to Submit the Assignment**

Before submitting, select `Kernel` -> `Restart Kernel & Run All Cells` if you are using a local environment, or `Runtime -> Restart session` and run all if using Google Colab, to ensure everything works as expected.

Next, create a submission folder with the following structure:
- Folder named `MSSV` (for example, if your student ID is `1234567`, name the folder `1234567`)
    - File `HW01.ipynb` (no need to submit other files)

Finally, compress this `MSSV` folder in `.zip` format (not `.rar` or any other format) and submit it via the link on Moodle.\
<font color=red>Please make sure to strictly follow this submission guideline.</font>

---
# 1. Import

```python
# Necessary Packages
import time
import requests
import numpy as np
import pandas as pd
from bs4 import BeautifulSoup

# YOUR CODE HERE (OPTION)
# If you need other support packages
import json
from tqdm import tqdm

```

---
# 2. Collect data from a website by parsing HTML (3 points)

In this section, you are tasked with collecting data from a website that simulates the sale of Pokémon. We have provided all the necessary URLs in a file called pokemon.txt, which contains a list of links you need to crawl for data. Each link corresponds to a page with details about different Pokémon for sale.

Your goal is to extract specific data from each of these pages and compile it into a structured format.

**Expected Output:** You need to create a `DataFrame` containing the following columns for each Pokémon:
- `SKU`: The unique identifier (ID) for each Pokémon.
- `Name`: The name of the Pokémon.
- `Price`: The sale price of the Pokémon.
- `InStock`: The quantity of this Pokémon currently available in stock.
- `Categories`: The category under which the Pokémon is listed (e.g., type, special edition).
- `Tags`: Any additional tags or attributes assigned to the Pokémon (e.g., rare, legendary, etc.).

**What You Need to Do**:
- Implement the `collect_data` function, which takes one input parameter: `course_urls_file`. This file contains all the URLs you need to process, each on a new line.
- For each URL in `pokemon.txt`, crawl the webpage and extract the relevant information for all the Pokémon listed on that page.
- The data you collect should be organized into a pandas `DataFrame` with the six specified columns (`SKU`, `Name`, `Price`, `InStock`, `Categories`, and `Tags`).
- The output DataFrame should resemble the structure of the provided sample file pokemon_example.csv. This sample contains a few examples to help you visualize the expected result format.

**Notes:**
- **Price Format**: Ensure the price is captured as a numerical value. If there are symbols like "$", remove them before storing the price.
- **In Stock**: If a Pokémon is out of stock, mark its quantity as 0 in the InStock field.
- **Categories & Tags**: Some Pokémon might belong to multiple categories or have multiple tags. Make sure to capture all relevant information and store them as lists or comma-separated strings.

```python
pokemon_example = pd.read_csv('./assets/pokemon_example.csv')
pokemon_example
```

Output:
```
SKU        Name  Price  InStock       Categories  \
0  4391   Bulbasaur   63.0       45    Pokemon, Seed   
1  7227     Ivysaur   87.0      142    Pokemon, Seed   
2  7036    Venusaur  105.0       30    Pokemon, Seed   
3  9086  Charmander   48.0      206  Lizard, Pokemon   
4  6565  Charmeleon  165.0      284   Flame, Pokemon   

                        Tags  
0  bulbasaur, Overgrow, Seed  
1    ivysaur, Overgrow, Seed  
2   Overgrow, Seed, venusaur  
3  Blaze, charmander, Lizard  
4   Blaze, charmeleon, Flame
```

<!-- grade_id: cell-d229a1a778f66947, points: 0, grade: false, solution: true -->
```python

def collect_data(course_urls_file: str) -> pd.DataFrame:
    """
    Collect data from a list of URLs provided in a file.

    Parameters
    ----------
    course_urls_file : str
        Path to the file containing URLs (e.g., 'Lab01/pokemon.txt').

    Returns
    -------
    pd.DataFrame
        A DataFrame containing the required data fields.
    """
    # Load URLs from file
    with open(course_urls_file, 'r') as file:
        urls = [line.strip() for line in file]

    # Initialize empty lists to store the values of each attribute
    sku, names, prices, in_stocks, categories, tags = [], [], [], [], [], []

    # YOUR CODE HERE

    def extract_product_information(url):
        response = requests.get(url)
        soup = BeautifulSoup(response.text, 'html.parser')
        
        summary_section = soup.find('div', class_='summary entry-summary')
        
        sku_element = summary_section.find('span', class_='sku')
        product_sku = int(sku_element.text) if sku_element else None
        
        name_element = summary_section.find('h1', class_='product_title entry-title')
        product_name = name_element.text if name_element else "Unknown Product"
        
        price_element = summary_section.find('span', class_='woocommerce-Price-amount amount')
        product_price = float(price_element.text[1:]) if price_element else 0.0
        
        stock_element = summary_section.find('p', class_='stock in-stock')
        stock_text = stock_element.text if stock_element else "0 in stock"
        product_stock = int(stock_text.replace(' in stock', '').strip()) if 'in stock' in stock_text else 0
        
        categories_element = summary_section.find('span', class_='posted_in')
        product_categories = ', '.join([cat.text for cat in categories_element.find_all('a')]) if categories_element else "No Categories"
        
        tags_element = summary_section.find('span', class_='tagged_as')
        product_tags = ', '.join([tag.text for tag in tags_element.find_all('a')]) if tags_element else "No Tags"
        
        return (product_sku, product_name, product_price, product_stock, product_categories, product_tags)

    craw_data = []
    for url in tqdm(urls):
        craw_data += [extract_product_information(url)]
        
    sku, names, prices, in_stocks, categories, tags = zip(*craw_data)
    # Create DataFrame with collected data
    data = pd.DataFrame({
        "SKU": sku,
        "Name": names,
        "Price": prices,
        "InStock": in_stocks,
        "Categories": categories,
        "Tags": tags
    })

    return data

```

<!-- grade_id: cell-7e3677f4a926ad6c, points: 3, grade: true, solution: false -->
```python
# TEST
data_pokemon = collect_data("./assets/pokemon.txt")
assert data_pokemon.shape == (755, 6), f"Expected shape (755, 6), got {data_pokemon.shape}"
```

Output:
```
100%|██████████| 755/755 [16:17<00:00,  1.29s/it]
```

```python
# Save to csv file with name pokemon.csv to grade
data_pokemon.to_csv("student_pokemon.csv", sep='\t', encoding='utf-8', index=False, header=True)
```

---
# 3. Collect data using Web API (4 points)

In this section, your task is to practice web data crawling using the World Bank API. You will be working with demographic and statistical data provided by the World Bank, covering population, education, health, and other key indicators for all countries from 1960 to 2022.

**Selected Indicators:** You will collect data for the following indicators
- `SP.POP.TOTL` – Total Population
- `SP.POP.TOTL.FE.IN` – Total Female Population
- `SP.POP.TOTL.MA.IN` – Total Male Population
- `SP.DYN.CBRT.IN` – Birth Rate
- `SP.DYN.CDRT.IN` – Death Rate
- `SP.DYN.LE00.MA.IN` – Average Life Expectancy (Male)
- `SP.DYN.LE00.FE.IN` – Average Life Expectancy (Female)
- `SE.PRM.ENRR` – Primary School Enrollment Rate
- `SE.TER.ENRR` – High School Enrollment Rate
- `SE.PRM.CMPT.ZS` – Primary School Completion Rate
- `SE.ADT.1524.LT.ZS` – Literacy Rate (Ages 15-24)

**Countries of Interest**: You are required to collect data from the following 7 countries as
- United States of America (US)
- India (IN)
- China (CN)
- Japan (JP)
- Canada (CA)
- Great Britain (GB)
- South Africa (ZA)

**Task Overview:** You are to implement a data collection function that queries the World Bank API for the specified indicators across these 7 countries. The data will be collected for each year from 1960 to 2022 and stored in a pandas `DataFrame` named `data_countries`.


You can expand your work on collecting data (such as collecting data from other countries and other indicators) by reading: https://datahelpdesk.worldbank.org/knowledgebase/articles/889392-api-documentation

**Hints**:

- Use the based URL: http://api.worldbank.org/v2/
- In order to collect data for each indicator of each country, you can use the URL: "http://api.worldbank.org/v2/countries/{country_code}/indicators/{indicator_code}"
    + `country_code` and `indicator_code` are provided above.
    + For example, you can use the following URL to get the `Total population` of Japan: http://api.worldbank.org/v2/countries/jp/indicators/SP.POP.TOTL

```python
data_countries_examples = pd.read_csv("./assets/countries_example.csv")
data_countries_examples
```

Output:
```
Total Population  Female Population  Male Population  Birth Rate  \
0       333287557.0        168266219.0      165021339.0         NaN   
1       332031554.0        167550001.0      164481553.0        11.0   
2       331511512.0        167203010.0      164308503.0        10.9   
3       328329953.0        165599805.0      162730147.0        11.4   
4       326838199.0        164926348.0      161911851.0        11.6   
5       325122128.0        164151818.0      160970309.0        11.8   
6       323071755.0        163224028.0      159847727.0        12.2   
7       320738994.0        162158414.0      158580581.0        12.4   
8       318386329.0        161084758.0      157301571.0        12.5   
9       316059947.0        160034189.0      156025758.0        12.4   

   Death Rate  Male life expectancy  Female life expectancy   \
0         NaN                   NaN                      NaN   
1      10.400                  73.5                     79.3   
2      10.300                
```

```python
BASE_URL = "https://api.worldbank.org/v2/"
COUNTRIES = ["US", "IN", "CN", "JP", "CA", "GB", "ZA"]
INDICATORS = [
    "SP.POP.TOTL",
    "SP.POP.TOTL.FE.IN",
    "SP.POP.TOTL.MA.IN",
    "SP.DYN.CBRT.IN",
    "SP.DYN.CDRT.IN",
    "SP.DYN.LE00.MA.IN",
    "SP.DYN.LE00.FE.IN",
    "SE.PRM.ENRR",
    "SE.TER.ENRR",
    "SE.PRM.CMPT.ZS",
    "SE.ADT.1524.LT.ZS",
]

# YOUR CODE HERE (option)
# If you need other initializations
countryMap = {
    "us": "United States",
    "in": "India",
    "cn": "China",
    "jp": "Japan",
    "ca": "Canada",
    "gb": "Great Britain",
    "za": "South Africa",
}

featureMap = {
    "SP.POP.TOTL": "Total Population",
    "SP.POP.TOTL.FE.IN": "Female Population",
    "SP.POP.TOTL.MA.IN": "Male Population",
    "SP.DYN.CBRT.IN": "Birth Rate",
    "SP.DYN.CDRT.IN": "Death Rate",
    "SP.DYN.LE00.MA.IN": "Life Expectancy Male",
    "SP.DYN.LE00.FE.IN": "Life Expectancy Female",
    "SE.PRM.ENRR": "Primary School Enrollment",
    "SE.TER.ENRR": "High School Enrollment",
    "SE.PRM.CMPT.ZS": "Primary Completion Rate",
    "SE.ADT.1524.LT.ZS": "Literacy rate"
}
```

<!-- grade_id: cell-d7f9f143f1ba0448, points: 0, grade: false, solution: true -->
```python
def collect_data(country_code: str, per_page: int, start_year: int, end_year: int, max_retries: int = 3, delay: int = 2) -> pd.DataFrame:
    """
    Collect data from the World Bank API for a specific country and date range.

    Parameters
    ----------
    country_code : str
        The ISO 3166-1 alpha-2 code for the country (e.g., 'US' for the United States).
    per_page : int
        Number of records to return per page.
    start_year : int
        The starting year of the data range.
    end_year : int
        The ending year of the data range.
    max_retries : int, optional
        Maximum number of retries for API requests in case of server errors (default is 3).
    delay : int, optional
        Delay (in seconds) between retries (default is 2).

    Returns
    -------
    pd.DataFrame
        A DataFrame containing the data collected from the API.

    Raises
    ------
    ValueError
        If the API request fails or if the data is not available.
    """
    # YOUR CODE HERE
    
    data = []
    
    # Get the country name using the country code
    country_name = countryMap.get(country_code.lower(), "Unknown Country")

    # Loop through each year in the specified range
    for year in range(start_year, end_year + 1):
        # Create a row for the current year and country
        row = {"Year": year, "Country": country_name}

        # Loop through each indicator to fetch its data
        for indicator in INDICATORS:
            url = f"http://api.worldbank.org/v2/country/{country_code}/indicator/{indicator}?format=json&date={year}"
            attempt = 0
            
            while attempt < max_retries:
                try:
                    # Make the API request
                    response = requests.get(url)
                    response.raise_for_status()  # Check for request errors
                    
                    json_data = response.json()
                    if len(json_data) > 1 and json_data[1]:
                        # Extract the value from the response
                        value = json_data[1][0].get('value', None)
                    else:
                        value = None
                    
                    # Store the value in the row
                    row[featureMap[indicator]] = value
                    break  # Exit retry loop if successful
                except requests.exceptions.RequestException:
                    attempt += 1
                    if attempt < max_retries:
                        print(f"Error fetching {indicator} for {year}. Retrying in {delay} seconds...")
                        time.sleep(delay)  # Wait before retrying
                    else:
                        print(f"Failed to fetch {indicator} for {year} after {max_retries} attempts.")
                        row[featureMap[indicator]] = None  # Assign None if failed after retries

        # Add the completed row to the data list
        data.append(row)
    
    # Check if any data was collected
    if not data:
        raise ValueError(f"No data available for country {country_code}")
    
    # Convert the data list into a DataFrame
    df = pd.DataFrame(data)
    return df

    
    
```

```python
def generate_countries_dataset(country_code_list: list) -> pd.DataFrame:
    """
    Generate a dataset by collecting data for a list of countries.

    Parameters
    ----------
    country_code_list : list of str
        List of ISO 3166-1 alpha-2 country codes (e.g., ['US', 'IN', 'CN']).

    Returns
    -------
    pd.DataFrame
        A DataFrame containing the combined data for all countries.

    Raises
    ------
    ValueError
        If data collection for any country fails.
    """
    # Initialize an empty list to store DataFrames
    data_frames = []

    for country_code in country_code_list:
        try:
            # Collect data for each country and append the result to the list
            df = collect_data(country_code=country_code,
                              per_page=100, start_year=2000, end_year=2022)
            data_frames.append(df)
        except ValueError as e:
            print(f"Error collecting data for {country_code}: {e}")

    # Concatenate all collected DataFrames into a single DataFrame
    combined_data = pd.concat(data_frames, ignore_index=True)

    return combined_data
```

<!-- grade_id: cell-f67d0a7fbf4ed980, points: 4, grade: true, solution: false -->
```python
# TEST
data_countries = generate_countries_dataset(COUNTRIES)
assert data_countries.shape == (161, 13)

```

```python
# Save to csv file with name coutries.csv to grade
data_countries.to_csv("student_countries.csv", sep='\t', encoding='utf-8', index=False, header=True)
```

---
# 4. Crawl data from Springer Journal (3 points)

As a Computer Science student conducting research on a specific topic, it's essential to read papers from academic journals.

In this assignment, you will work with the journal SN Computer Science. Your task is to crawl and extract detailed information for every paper published in this journal. You can find an overview of the journal on the [main page](https://link.springer.com/journal/42979), which states:
```
SN Computer Science is a broad-based, hybrid, peer reviewed journal that publishes original research in all the disciplines of computer science including various inter-disciplinary aspects. The journal aims to be a global forum of, for, and by the community and offers:
```

Using your previous code as a foundation, you need to create a table (dataframe) that includes the following information for each paper:
- `Title`: The title of the paper
- `Date`: The publication date
- `Link`: The URL to the paper
- `Authors`: The authors of the paper
- `Affiliations`: The affiliations of each author


When you are a Computer Science student and going to do research on a specific topic, you need to read paper from journal. 

In this task, I will give you a journal, SN Computer Science and you need to crawl all information of each paper which published in this journal. You may find an overview of this journal in the  as:


Here is an example of the expected output (dataframe):

```python
data_journal_examples = pd.read_csv("./assets/journal_example.csv")
data_journal_examples
```

Output:
```
Title               Date  \
0  An Upgraded Approach for Identifying Partially...  20 September 2024   
1  A Comprehensive Review on Deep Learning Techni...  20 September 2024   
2  Hybrid Deep Learning Approach with Feature Eng...  20 September 2024   
3  Roberta and BERT: Revolutionizing Mental Healt...  19 September 2024   
4  An Intelligent Image Encryption Scheme Based o...  17 September 2024   
5  A Systematic Review on Federated Learning in E...  17 September 2024   
6  Park-Net: A Deep Model for Early Detection of ...  17 September 2024   
7  Leveraging Deep Embedding Models for Arabic Te...  17 September 2024   
8  Bioanalytical Method Development and Validatio...  17 September 2024   

                                                Link  \
0  https://link.springer.com/article/10.1007/s429...   
1  https://link.springer.com/article/10.1007/s429...   
2  https://link.springer.com/article/10.1007/s429...   
3  https://link.springer.com/article/10.1007/s429...   
4  https://link.
```

Complete this task by writing code that scrapes and organizes this information into a structured format.

<!-- grade_id: cell-2c2e1708d3eb97c2, points: 3, grade: true, solution: true -->
```python
def scrape_journal_first_n_page(n):
    total_pages = n
    all_articles = []

    for page in range(1, total_pages + 1):
        BASE_URL = f'https://link.springer.com/journal/42979/articles?filterOpenAccess=false&page={page}'
        response = requests.get(BASE_URL)
        soup = BeautifulSoup(response.text, 'html.parser')

        for article in soup.find_all('article', class_='app-card-open app-card-open--has-image'):
            title = article.find('h3', class_="app-card-open__heading").get_text(strip=True)
            date = get_date(article)
            link = get_link(article)

            new_response = requests.get(link)
            new_soup = BeautifulSoup(new_response.text, 'html.parser')
            authors = get_author(new_soup)
            affiliations = get_affiliations(new_soup)

            all_articles.append([title, date, link, authors, affiliations])

    df = pd.DataFrame(all_articles, columns=['Title', 'Date', 'Link', 'Authors', 'Affiliations'])
    df.to_csv("student_journal.csv", sep='\t', encoding='utf-8', index=False, header=True)

def get_date(article):
    months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', 'October', 'November', 'December']
    meta_div = article.find('div', class_='app-card-open__meta')
    if meta_div:
        tags = meta_div.find_all('span', class_='c-meta__item')
        for tag in tags:
            text = tag.get_text(strip=True)
            if "Open access" in text:
                continue
            if any(month in text for month in months):
                return text
    return None

def get_link(article):
    article_link = article.find('h3', class_="app-card-open__heading").find('a')
    return article_link['href'] if article_link else None

def get_author(new_soup):
    authors = new_soup.find_all('li', class_='c-article-author-list__item')
    author_names = [author.find('a').get_text(strip=True) for author in authors if author.find('a').get_text(strip=True) not in ["…", ""]]
    return ', '.join(author_names[:-1]) + " & " + author_names[-1] if len(author_names) > 1 else author_names[0] if author_names else None

def get_affiliations(new_soup):
    affiliations = new_soup.find_all('li', id=lambda x: x and x.startswith('Aff'))
    return "; ".join(f"{aff.find('p', class_='c-article-author-affiliation__address').get_text(strip=True)}: {aff.find('p', class_='c-article-author-affiliation__authors-list').get_text(strip=True)}" for aff in affiliations) if affiliations else None


scrape_journal_first_n_page(1)
student_journal = pd.read_csv("./student_journal.csv", sep="\t")
student_journal
```

Output:
```
Title             Date  \
0   An Efficient Approach for VM and Database Segm...  19 October 2024   
1   Metaphor Identification and Interpretation in ...  18 October 2024   
2   Inclusive Learning and Assessment in the Era o...  18 October 2024   
3   Machine Learning Based Assessment of Elite Foo...  18 October 2024   
4   Optical Sensor Based Continuous Blood Glucose ...  18 October 2024   
5   Drone-Based Intelligent System for Social Dist...  17 October 2024   
6   Events Correlations for Fault Identification i...  17 October 2024   
7   User Task Priority Based Resource Allocation w...  17 October 2024   
8         Smart Agriculture: A Comprehensive Overview  17 October 2024   
9   Assessing Player Contributions in League of Le...  16 October 2024   
10  Cytosine Methylation and Demethylation Model: ...  15 October 2024   
11  Early Detection of Benign Ovarian Tumor Classi...  15 October 2024   
12  Exploring Advanced Deep Learning Paradigms for...  15 October 2024   
13  Lie Comm
```

```python

```